rule_files:
- kube-state-metrics.rules.yaml

evaluation_interval: 30s

tests:
- interval: 30s
  input_series:
  # KubeStateMetricsShootDown
  - series: 'up{job="kube-state-metrics", type="shoot"}'
    values: '0+0x30'
  # KubeStateMetricsSeedDown
  - series: 'kube_verticalpodautoscaler_status_recommendation_containerrecommendations_target{exported_job="kube-state-metrics"}'
    values: '42 _ stale'
  # NoWorkerNodes
  - series: 'kube_node_spec_unschedulable'
    values: '2+0x20'
  - series: 'kube_node_spec_unschedulable'
    values: '1+0x50'
  - series: 'kube_node_info'
    values: '1+0x50'
  alert_rule_test:
  - eval_time: 15m
    alertname: KubeStateMetricsShootDown
    exp_alerts:
    - exp_labels:
        service: kube-state-metrics-shoot
        severity: info
        visibility: operator
        type: seed
      exp_annotations:
        summary: Kube-state-metrics for shoot cluster metrics is down.
        description: There are no running kube-state-metric pods for the shoot cluster. No kubernetes resource metrics can be scraped.
  - eval_time: 20m
    alertname: KubeStateMetricsSeedDown
    exp_alerts:
    - exp_labels:
        service: kube-state-metrics-seed
        severity: critical
        visibility: operator
        type: seed
      exp_annotations:
        summary: There are no kube-state-metrics metrics for the control plane
        description: Kube-state-metrics is scraped by the cache prometheus and federated by the control plane prometheus. Something is broken in that process.
  - eval_time: 25m
    alertname: NoWorkerNodes
    exp_alerts:
    - exp_labels:
        service: nodes
        severity: blocker
        visibility: all
      exp_annotations:
        summary: No nodes available. Possibly all workloads down.
        description: There are no worker nodes in the cluster or all of the worker nodes in the cluster are not schedulable.
