rule_files:
- kube-etcd3-main.rules.yaml

evaluation_interval: 30s

tests:
- interval: 30s
  input_series:
  # KubeEtcdMainDown
  # KubeEtcdBackupRestoreMainDown
  - series: 'up{job="kube-etcd3-main"}'
    values: '1+0x70'
  # KubeEtcd3MainNoLeader
  - series: 'etcd_server_has_leader{job="kube-etcd3-main"}'
    values: '0+0x30'
  # KubeEtcd3HighNumberOfFailedProposals
  - series: 'etcd_server_proposals_failed_total{job="kube-etcd3-main", pod="etcd"}'
    values: '0+1x6 6+0x115'
  # KubeEtcd3DbSizeLimitApproaching
  # KubeEtcd3DbSizeLimitCrossed
  - series: 'etcd_mvcc_db_total_size_in_bytes{job="kube-etcd3-main"}'
    values: '7194070000+107374182x20' # 6.7GB 6.8GB 6.9GB .. 7.7GB
  # KubeEtcdDeltaBackupFailed
  - series: 'etcdbr_snapshot_latest_timestamp{job="kube-etcd3-backup-restore-main",kind="Incr", pod="etcd", role="test"}'
    values: '0+0x62'
  - series: 'etcdbr_snapshot_required{job="kube-etcd3-backup-restore-main",kind="Incr", pod="etcd", role="test"}'
    values: '1+0x62'
  - series: 'etcd_server_is_leader{job="kube-etcd3-main", pod="etcd", role="test"}'
    values: '1+0x70'
  # KubeEtcdFullBackupFailed
  - series: 'etcdbr_snapshot_latest_timestamp{job="kube-etcd3-backup-restore-main",kind="Full", pod="etcd", role="test"}'
    values: '0+0x2912'
  - series: 'etcdbr_snapshot_required{job="kube-etcd3-backup-restore-main",kind="Full", pod="etcd", role="test"}'
    values: '1+0x2912'
  - series: 'etcd_server_is_leader{job="kube-etcd3-main", pod="etcd", role="test"}'
    values: '1+0x2912'
  # KubeEtcdRestorationFailed
  - series: 'etcdbr_restoration_duration_seconds_count{job="kube-etcd3-backup-restore-main",succeeded="false"}'
    values: '0+0x7 1 2 2'
  # KubeEtcdBackupRestoreMainDown
  - series: 'up{job="kube-etcd3-backup-restore-main"}'
    values: '0+0x70'
  - series: 'etcdbr_snapshotter_failure{job="kube-etcd3-backup-restore-main"}'
    values: '1+1x30 1+0x40'
  alert_rule_test:
  - eval_time: 15m
    alertname: KubeEtcdMainDown
    exp_alerts:
    - exp_labels:
        service: etcd
        severity: critical
        type: seed
        visibility: operator
      exp_annotations:
        description: Etcd3 cluster main is unavailable (due to possible quorum loss) or cannot be scraped. As long as etcd3 main is down, the cluster is unreachable.
        summary: Etcd3 main cluster down.
  - eval_time: 15m
    alertname: KubeEtcd3MainNoLeader
    exp_alerts:
    - exp_labels:
        service: etcd
        severity: critical
        type: seed
        visibility: operator
      exp_annotations:
        description: Etcd3 main has no leader. Possible network partition in the etcd cluster.
        summary: Etcd3 main has no leader.
  - eval_time: 1h
    alertname: KubeEtcd3HighNumberOfFailedProposals
    exp_alerts:
    - exp_labels:
        service: etcd
        severity: warning
        type: seed
        visibility: operator
        pod: etcd
        job: kube-etcd3-main
      exp_annotations:
        description: Etcd3 main pod etcd has seen 6 proposal failures within the last hour.
        summary: High number of failed etcd proposals
  - eval_time: 5m
    alertname: KubeEtcd3DbSizeLimitApproaching
    exp_alerts:
    - exp_labels:
        job: kube-etcd3-main
        service: etcd
        severity: warning
        type: seed
        visibility: all
      exp_annotations:
        description: Etcd3 main DB size is approaching its current practical limit of 8GB. Etcd quota might need to be increased.
        summary: Etcd3 main DB size is approaching its current practical limit.
  - eval_time: 10m
    alertname: KubeEtcd3DbSizeLimitCrossed
    exp_alerts:
    - exp_labels:
        job: kube-etcd3-main
        service: etcd
        severity: critical
        type: seed
        visibility: all
      exp_annotations:
        description: Etcd3 main DB size has crossed its current practical limit of 8GB. Etcd quota must be increased to allow updates.
        summary: Etcd3 main DB size has crossed its current practical limit.
  - eval_time: 31m
    alertname: KubeEtcdDeltaBackupFailed
    exp_alerts:
    - exp_labels:
        pod: etcd
        role: test
        service: etcd
        severity: critical
        type: seed
        visibility: operator
      exp_annotations:
        description: No delta snapshot for the past 30 minutes have been taken by backup-restore leader.
        summary: Etcd delta snapshot failure.
  - eval_time: 1456m
    alertname: KubeEtcdFullBackupFailed
    exp_alerts:
    - exp_labels:
        pod: etcd
        role: test
        service: etcd
        severity: critical
        type: seed
        visibility: operator
      exp_annotations:
        description: No full snapshot for at least last 24 hours have been taken by backup-restore leader.
        summary: Etcd full snapshot failure.
  - eval_time: 5m
    alertname: KubeEtcdRestorationFailed
    exp_alerts:
    - exp_labels:
        job: kube-etcd3-backup-restore-main
        succeeded: false
        service: etcd
        severity: critical
        type: seed
        visibility: operator
      exp_annotations:
        description: Etcd data restoration was triggered, but has failed.
        summary: Etcd data restoration failure.
  - eval_time: 30m
    alertname: KubeEtcdBackupRestoreMainDown
    exp_alerts:
    - exp_labels:
        service: etcd
        severity: critical
        type: seed
        visibility: operator
      exp_annotations:
        description: Etcd backup restore main process down or snapshotter failed with error. Backups will not be triggered unless backup restore is brought back up. This is unsafe behaviour and may cause data loss.
        summary: Etcd backup restore main process down or snapshotter failed with error
